

/*---
esid: sec-regexp.prototype-@@match
description: RegExp.prototype[@@match] behavior with 'v' flag, comparing with 'u' flag
features: [Symbol.match, regexp-v-flag]
includes: [compareArray.js]
---*/

const text = 'р ®·aр ®·bр ®·cрџ‘ЁвЂЌрџ‘©вЂЌрџ‘§вЂЌрџ‘¦d';

function doMatch(regex) {
  return RegExp.prototype[Symbol.match].call(regex, text);
}

assert.compareArray(doMatch(/р ®·/g), ["р ®·", "р ®·", "р ®·"], "Basic match with g flag");
assert.compareArray(doMatch(/р ®·/u), ["р ®·"], "Match with u flag");
assert.compareArray(doMatch(/р ®·/v), ["р ®·"], "Match with v flag");

assert.compareArray(doMatch(/\p{Script=Han}/gu), ["р ®·", "р ®·", "р ®·"], "Unicode property escapes with u flag");
assert.compareArray(doMatch(/\p{Script=Han}/gv), ["р ®·", "р ®·", "р ®·"], "Unicode property escapes with v flag");

assert.compareArray(doMatch(/./g), ["\uD842", "\uDFB7", "a", "\uD842", "\uDFB7", "b", "\uD842", "\uDFB7", "c", "\uD83D", "\uDC68", "\u200D", "\uD83D", "\uDC69", "\u200D", "\uD83D", "\uDC67", "\u200D", "\uD83D", "\uDC66", "d"], "Dot without u or v flag");
assert.compareArray(doMatch(/./gu), ["р ®·", "a", "р ®·", "b", "р ®·", "c", "рџ‘Ё", "вЂЌ", "рџ‘©", "вЂЌ", "рџ‘§", "вЂЌ", "рџ‘¦", "d"], "Dot with u flag");
assert.compareArray(doMatch(/./gv), ["р ®·", "a", "р ®·", "b", "р ®·", "c", "рџ‘Ё", "вЂЌ", "рџ‘©", "вЂЌ", "рџ‘§", "вЂЌ", "рџ‘¦", "d"], "Dot with v flag");

assert.compareArray(doMatch(/[рџ‘ЁвЂЌрџ‘©вЂЌрџ‘§вЂЌрџ‘¦]/v), ["рџ‘Ё"], "Complex emoji sequence in set notation with v flag");
assert.compareArray(doMatch(/[рџ‘ЁвЂЌрџ‘©вЂЌрџ‘§вЂЌрџ‘¦]/u), ["рџ‘Ё"], "Complex emoji sequence in set notation with u flag throws");

assert.sameValue(doMatch(/x/u), null, "Non-matching regex with u flag");
assert.sameValue(doMatch(/x/v), null, "Non-matching regex with v flag");

