

/*---
esid: sec-temporal.duration.from
description: >
  Throws when a fractional unit is present and a sub-part is zero.
features: [Temporal]
---*/

const invalid = [
  
  "PT0.1H0M",

  
  "PT0.1H0.0M",

  
  "PT0.1H0S",

  
  "PT0.1H0.0S",

  
  "PT0.1M0S",

  
  "PT0.1M0.0S",
];

for (let string of invalid) {
  assert.throws(RangeError, () => Temporal.Duration.from(string));
}
