

/*---
esid: sec-regexp.prototype-@@replace
description: RegExp.prototype[@@replace] behavior with 'v' flag
features: [Symbol.replace, regexp-v-flag]
---*/

const text = 'đ ®·ađ ®·bđ ®·';

function doReplace(regex, replacement) {
  return RegExp.prototype[Symbol.replace].call(regex, text, replacement);
}

assert.sameValue(doReplace(/đ ®·/g, '-'), "-a-b-", "Basic replace with g flag");
assert.sameValue(doReplace(/đ ®·/v, '-'), "-ađ ®·bđ ®·", "Replace with v flag");
assert.sameValue(doReplace(/\p{Script=Han}/gv, 'X'), "XaXbX", "Unicode property escapes with v flag");
assert.sameValue(doReplace(/./gv, '$&$&'), "đ ®·đ ®·aađ ®·đ ®·bbđ ®·đ ®·", "Dot with v flag");
assert.sameValue(
  doReplace(/./gv, (match, index) => `[${match}:${index}]`),
  "[đ ®·:0][a:2][đ ®·:3][b:5][đ ®·:6]",
  "Replace with function"
);

assert.sameValue(doReplace(/\p{Script=Han}/gu, 'X'), "XaXbX", "Unicode property escapes with u flag");
assert.sameValue(doReplace(/\p{Script=Han}/gv, 'X'), "XaXbX", "Unicode property escapes with v flag");
