

/*---
esid: sec-temporal.zoneddatetime.prototype.round
description: String argument creates options object with null prototype, so Object.prototype pollution doesn't affect it
features: [Temporal]
---*/

const props = ["roundingIncrement", "roundingMode"];
for (const prop of props) {
  Object.defineProperty(Object.prototype, prop, {
    get() {
      throw new Test262Error(`Object.prototype.${prop} was looked up`);
    },
    configurable: true,
  });
}

try {
  const instance = new Temporal.ZonedDateTime(1_000_000_000_123_456_789n, "UTC");
  instance.round("hour");
} finally {
  for (const prop of props) {
    delete Object.prototype[prop];
  }
}
