

/*---
esid: sec-names-and-keywords
description: Test grandfathered characters of ID_Continue.
info: |
  Grandfathered characters (Other_ID_Start + Other_ID_Continue)
---*/


var a℘; 
var a℮; 
var a゛; 
var a゜; 


var aᢅ; 
var aᢆ; 


var a፩; 
var a፪; 
var a፫; 
var a፬; 
var a፭; 
var a፮; 
var a፯; 
var a፰; 
var a፱; 


var a·; 
var a·; 


var a᧚; 
