

/*---
esid: sec-array.prototype.findlastindex
description: >
  Predicate is called for each array property.
info: |
  Array.prototype.findLastIndex ( predicate[ , thisArg ] )

  ...
  4. Let k be len - 1.
  5. Repeat, while k ≥ 0,
    ...
    c. Let testResult be ! ToBoolean(? Call(predicate, thisArg, « kValue, 𝔽(k), O »)).
  ...
features: [array-find-from-last]
---*/

var arr = [undefined, , , 'foo'];
var called = 0;

arr.findLastIndex(function() {
  called++;
});

assert.sameValue(called, 4);
