

/*---
esid: sec-create-bytes-module
description: Creates bytes module from txt file
flags: [module]
features: [import-attributes, immutable-arraybuffer, import-bytes]
includes: [compareArray.js]
---*/

import value from './bytes-from-txt_FIXTURE.txt' with { type: 'bytes' };

assert(value instanceof Uint8Array);
assert(value.buffer instanceof ArrayBuffer);

assert.sameValue(value.length, 13);
assert.sameValue(value.buffer.byteLength, 13);
assert.sameValue(value.buffer.immutable, true);

assert.compareArray(
  Array.from(value),
  [
    0x48, 
    0x65, 
    0x6c, 
    0x6c, 
    0x6f, 
    0x20, 
    0x57, 
    0x6f, 
    0x72, 
    0x6c, 
    0x64, 
    0x21, 
    0x0a, 
  ]
);

assert.throws(TypeError, function() {
  value.buffer.resize(0);
});

assert.throws(TypeError, function() {
  value.buffer.transfer();
});
